---
title       : Insert the chapter title here
description : Insert the chapter description here


--- type:NormalExercise lang:r xp:100 skills:1 key:37dde722b6
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}
library(OneR)
OneR(iris)
```

*** =sample_code
```{r}
library(OneR)
OneR(iris)
```

*** =solution
```{r}
library(OneR)
OneR(iris)
```

*** =sct
```{r}

```
